[1506337976] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506338036] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506338089] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1506338096] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1506338126] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506338149] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1506338156] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506338186] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506338209] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1506338246] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506338269] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1506338306] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506339169] SERVICE ALERT: localhost;SSH;OK;HARD;4;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1506391199] SERVICE ALERT: uat-nginx;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506391259] SERVICE ALERT: uat-nginx;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506391319] SERVICE ALERT: uat-nginx;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506391379] SERVICE ALERT: uat-nginx;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506393182] SERVICE ALERT: uat-nginx;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506393242] SERVICE ALERT: uat-nginx;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506393302] SERVICE ALERT: uat-nginx;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506393362] SERVICE ALERT: uat-nginx;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506396249] SERVICE ALERT: nsca_server;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506398236] SERVICE ALERT: nsca_server;nsca;CRITICAL;SOFT;1;testOK
[1506398666] SERVICE ALERT: nsca_server;nsca;CRITICAL;SOFT;2;testOK
[1506399739] SERVICE ALERT: nsca_server;nsca;OK;SOFT;3;testOK:10
[1506415029] SERVICE ALERT: uat-database;HTTP;CRITICAL;SOFT;1;connect to address 10.0.101.10 and port 80: Connection refused
[1506415039] SERVICE ALERT: uat-services;HTTP;CRITICAL;SOFT;1;connect to address 10.0.100.10 and port 80: Connection refused
[1506415089] SERVICE ALERT: uat-database;HTTP;CRITICAL;SOFT;2;connect to address 10.0.101.10 and port 80: Connection refused
[1506415099] SERVICE ALERT: uat-services;HTTP;CRITICAL;SOFT;2;connect to address 10.0.100.10 and port 80: Connection refused
[1506415149] SERVICE ALERT: uat-database;HTTP;CRITICAL;SOFT;3;connect to address 10.0.101.10 and port 80: Connection refused
[1506415159] SERVICE ALERT: uat-services;HTTP;CRITICAL;SOFT;3;connect to address 10.0.100.10 and port 80: Connection refused
[1506415170] SERVICE ALERT: uat-database;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506415179] SERVICE ALERT: uat-services;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506415209] SERVICE ALERT: uat-database;HTTP;CRITICAL;HARD;4;connect to address 10.0.101.10 and port 80: Connection refused
[1506415219] SERVICE ALERT: uat-services;HTTP;CRITICAL;HARD;4;connect to address 10.0.100.10 and port 80: Connection refused
[1506415230] SERVICE ALERT: uat-database;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506415239] SERVICE ALERT: uat-services;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506415290] SERVICE ALERT: uat-database;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506415299] SERVICE ALERT: uat-services;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506415350] SERVICE ALERT: uat-database;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506415359] SERVICE ALERT: uat-services;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416344] SERVICE ALERT: uat-bastion;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506416404] SERVICE ALERT: uat-bastion;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506416438] SERVICE ALERT: uat-bastion;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416464] SERVICE ALERT: uat-bastion;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1506416479] SERVICE ALERT: uat-bastion;SSH;CRITICAL;SOFT;1;connect to address 10.0.0.10 and port 22: Connection refused
[1506416498] SERVICE ALERT: uat-bastion;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416524] SERVICE ALERT: uat-bastion;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506416539] SERVICE ALERT: uat-bastion;SSH;CRITICAL;SOFT;2;connect to address 10.0.0.10 and port 22: Connection refused
[1506416558] SERVICE ALERT: uat-bastion;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416599] SERVICE ALERT: uat-bastion;SSH;CRITICAL;SOFT;3;connect to address 10.0.0.10 and port 22: Connection refused
[1506416618] SERVICE ALERT: uat-bastion;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416659] SERVICE ALERT: uat-bastion;SSH;CRITICAL;HARD;4;connect to address 10.0.0.10 and port 22: Connection refused
[1506418975] HOST ALERT: dev-bastion;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1506419007] HOST ALERT: dev-database;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1506419020] HOST ALERT: dev-bastion;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1506419040] HOST ALERT: dev-nginx;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1506419063] HOST ALERT: dev-database;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1506419067] HOST ALERT: dev-bastion;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1506419074] HOST ALERT: dev-services;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1506419077] HOST ALERT: dev-nginx;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1506419094] SERVICE ALERT: dev-bastion;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506419099] SERVICE ALERT: dev-database;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506419104] SERVICE ALERT: dev-nginx;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506419108] SERVICE ALERT: dev-services;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506419124] HOST ALERT: dev-bastion;DOWN;SOFT;4;(Host check timed out after 30.02 seconds)
[1506419129] HOST ALERT: dev-database;DOWN;SOFT;3;(Host check timed out after 30.02 seconds)
[1506419134] HOST ALERT: dev-nginx;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1506419138] HOST ALERT: dev-services;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1506419142] SERVICE ALERT: dev-bastion;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1506419147] SERVICE ALERT: dev-database;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1506419151] SERVICE ALERT: dev-nginx;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1506419156] SERVICE ALERT: dev-services;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1506419172] HOST ALERT: dev-bastion;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1506419177] HOST ALERT: dev-database;DOWN;SOFT;4;(Host check timed out after 30.02 seconds)
[1506419181] HOST ALERT: dev-nginx;DOWN;SOFT;4;(Host check timed out after 30.03 seconds)
[1506419186] HOST ALERT: dev-services;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1506419208] HOST ALERT: dev-bastion;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1506419213] HOST ALERT: dev-database;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1506419217] HOST ALERT: dev-nginx;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1506419222] HOST ALERT: dev-services;DOWN;SOFT;4;(Host check timed out after 30.02 seconds)
[1506419235] SERVICE ALERT: dev-bastion;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506419239] SERVICE ALERT: dev-database;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506419244] SERVICE ALERT: dev-nginx;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506419249] SERVICE ALERT: dev-services;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506419263] HOST ALERT: dev-bastion;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1506419269] HOST ALERT: dev-database;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1506419272] HOST ALERT: dev-nginx;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1506419278] HOST ALERT: dev-services;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1506420345] SERVICE ALERT: dev-nginx;SSH;OK;SOFT;2;test
[1506420379] SERVICE ALERT: dev-bastion;HTTP;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506420379] SERVICE ALERT: dev-database;HTTP;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506420379] SERVICE ALERT: dev-nginx;HTTP;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506420379] SERVICE ALERT: dev-services;HTTP;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506420440] SERVICE ALERT: dev-services;PING;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1506420440] SERVICE ALERT: dev-database;PING;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1506420440] SERVICE ALERT: dev-nginx;PING;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1506420440] SERVICE ALERT: dev-bastion;PING;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1506420499] SERVICE ALERT: dev-bastion;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506421456] HOST ALERT: dev-bastion;UP;HARD;8;testOK
[1506421485] HOST ALERT: dev-nginx;UP;HARD;6;testOK
[1506421585] SERVICE ALERT: dev-nginx;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506421705] SERVICE ALERT: dev-services;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506421705] SERVICE ALERT: dev-database;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506431650] HOST ALERT: dev-database;UP;HARD;7;test
[1506431673] HOST ALERT: dev-services;UP;HARD;5;test
[1506437776] HOST ALERT: dev-bastion;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506437776] HOST ALERT: dev-nginx;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506437776] HOST ALERT: dev-database;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506437776] HOST ALERT: nsca-server;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506437776] HOST ALERT: dev-services;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506437937] SERVICE ALERT: dev-services;SSH;OK;HARD;3;SSH OK - OpenSSH_6.6.1 (protocol 2.0)
[1506437943] HOST ALERT: dev-nginx;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.03 ms
[1506437965] HOST ALERT: dev-database;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 2.74 ms
[1506437980] HOST ALERT: dev-services;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1506437982] SERVICE ALERT: dev-bastion;HTTP;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1506437983] HOST ALERT: dev-bastion;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1506438118] SERVICE ALERT: dev-nginx;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 3934 bytes in 0.006 second response time
[1506438151] SERVICE ALERT: dev-database;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 2.81 ms
[1506438160] SERVICE ALERT: dev-nginx;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 2.88 ms
[1506438406] HOST ALERT: nsca-server;UP;SOFT;2;OK
[1506438486] SERVICE ALERT: dev-nginx;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506438675] SERVICE ALERT: uat-bastion;SSH;OK;HARD;4;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1506438681] SERVICE ALERT: dev-bastion;HTTP;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1506438681] SERVICE ALERT: dev-services;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506438731] HOST ALERT: dev-bastion;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506438731] HOST ALERT: dev-services;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506438741] SERVICE ALERT: dev-nginx;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1506438791] HOST ALERT: dev-database;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506438791] HOST ALERT: dev-nginx;DOWN;SOFT;1;CRITICAL: Host State is Stale.
[1506438802] SERVICE ALERT: dev-database;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1506438802] SERVICE ALERT: dev-nginx;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out